# Stökiometribestämning

Stökiometri är ett viktigt begrepp inom kemin som handlar om de kvantitativa förhållandena mellan reaktanter och produkter i en kemisk reaktion. I denna övning kommer vi att skapa ett program som utför stökiometribestämningar.

Man börjar med en lexikon som heter "reaktion" och representerar en kemisk reaktion. Nycklarna i dictionaryn ska vara de kemiska ämnena (molekyler eller atomer) som är involverade i reaktionen, och värdena ska vara deras motsvarande stökiometriska koefficienter. 

**Uppgift 1:** Skapa en funktion som heter "stökiometri" som tar reaktionsdictionaryn och en given mängd (i mol) av en reaktant som inmatning. Funktionen ska beräkna mängden av alla andra ämnen som är involverade i reaktionen. Utdata ska vara en dictionary som innehåller de beräknade mängderna.

<!-- grade_id: cell-77149312d7328571, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-84580d6f75c41f33, points: 0, grade: true, solution: false -->
```python
# Reaktion 1: 2H2 + O2 -> 2H2O
reaktion1 = {"H2":2,"O2":1, "H2O":2}
assert stökiometri(reaktion1, "H2", 1) == {"H2":1,"O2":0.5, "H2O":1}
assert stökiometri(reaktion1, "O2", 3.2) == {"H2":6.4,"O2":3.2, "H2O":6.4}

# CH4 + 2O2 -> CO2 + 2H2O
reaktion2 = {"CH4":1,"O2":2, "CO2":1, "H2O":2}
assert stökiometri(reaktion2, "O2", 0.8) == {"CH4":0.4,"O2":0.8, "CO2":0.4, "H2O":0.8}

# 6CO2 + 6H2O -> C6H12O6 + 6O2
reaktion3 = {"CO2":6,"H2O":6, "C6H12O6":1, "O2":6}
assert stökiometri(reaktion3, "CO2", 0.9) == {"CO2":0.9,"H2O":0.9, "C6H12O6":0.15, "O2":0.9}
```

**Uppgift 2:** Modifiera det tidigare programmet för att istället ange en dictionary som inmatning med mängderna av varje reaktant. Om mängden av ett ämne inte anges, antag 0. Hitta den begränsande reaktanten och de slutliga mängderna av alla reaktanter och produkter som finns kvar i slutet av reaktionen, under antagande av 100% utbyte. Skriv ut den begränsande reaktanten. Utdata ska vara en dictionary som innehåller de beräknade mängderna.

Notera att nu har reaktion lexikon negativ tecken på produkter.

<!-- grade_id: cell-79009619aef9483a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-27761b87b7ef1aad, points: 0, grade: true, solution: false -->
```python
# Reaktion 1: 2H2 + O2 -> 2H2O
reaktion1 = {"H2":2,"O2":1, "H2O":-2}
assert stökiometri(reaktion1, {"H2":1, "O2":1}) == {'H2': 0.0, 'O2': 0.5, 'H2O': 1.0}
```

<!-- grade_id: cell-d0d9cf2a46425cac, points: 0, grade: true, solution: false -->
```python
# CH4 + 2O2 -> CO2 + 2H2O
reaktion2 = {"CH4":1,"O2":2, "CO2":-1, "H2O":-2}
assert stökiometri(reaktion2, {"CH4":0.8, "O2":5.0, "H2O":10}) == {"CH4":0,"O2":3.4, "CO2":0.8, "H2O":11.6}
```

<!-- grade_id: cell-87c77110b7beb5e3, points: 0, grade: true, solution: false -->
```python
# 6CO2 + 6H2O -> C6H12O6 + 6O2
reaktion3 = {"CO2":6,"H2O":6, "C6H12O6":-1, "O2":-6}
assert stökiometri(reaktion3, {"CO2":0.9}) == {"CO2":0.9,"H2O":0, "C6H12O6":0, "O2":0}
```

# Molekylmassa

**Uppgift:** Inspirerad av övningen "unik_atomlista", skapa en funktion som läser in en sträng och returnerar molekylmassan genom att använda den tillhandahållna ordboken för molmassa:

<!-- grade_id: cell-03070e6b6af15a5c, points: 0, grade: false, solution: true -->
```python
molmassa = {'H' : 1.008,'He' : 4.003, 'Li' : 6.941, 'Be' : 9.012,\
             'B' : 10.811, 'C' : 12.011, 'N' : 14.007, 'O' : 15.999,\
             'F' : 18.998, 'Ne' : 20.180, 'Na' : 22.990, 'Mg' : 24.305,\
             'Al' : 26.982, 'Si' : 28.086, 'P' : 30.974, 'S' : 32.066,\
             'Cl' : 35.453, 'Ar' : 39.948, 'K' : 39.098, 'Ca' : 40.078,\
             'Sc' : 44.956, 'Ti' : 47.867, 'V' : 50.942, 'Cr' : 51.996,\
             'Mn' : 54.938, 'Fe' : 55.845, 'Co' : 58.933, 'Ni' : 58.693,\
             'Cu' : 63.546, 'Zn' : 65.38, 'Ga' : 69.723, 'Ge' : 72.631,\
             'As' : 74.922, 'Se' : 78.971, 'Br' : 79.904, 'Kr' : 84.798,\
             'Rb' : 84.468, 'Sr' : 87.62, 'Y' : 88.906, 'Zr' : 91.224,\
             'Nb' : 92.906, 'Mo' : 95.95, 'Tc' : 98.907, 'Ru' : 101.07,\
             'Rh' : 102.906, 'Pd' : 106.42, 'Ag' : 107.868, 'Cd' : 112.414,\
             'In' : 114.818, 'Sn' : 118.711, 'Sb' : 121.760, 'Te' : 126.7,\
             'I' : 126.904, 'Xe' : 131.294, 'Cs' : 132.905, 'Ba' : 137.328,\
             'La' : 138.905, 'Ce' : 140.116, 'Pr' : 140.908, 'Nd' : 144.243,\
             'Pm' : 144.913, 'Sm' : 150.36, 'Eu' : 151.964, 'Gd' : 157.25,\
             'Tb' : 158.925, 'Dy': 162.500, 'Ho' : 164.930, 'Er' : 167.259,\
             'Tm' : 168.934, 'Yb' : 173.055, 'Lu' : 174.967, 'Hf' : 178.49,\
             'Ta' : 180.948, 'W' : 183.84, 'Re' : 186.207, 'Os' : 190.23,\
             'Ir' : 192.217, 'Pt' : 195.085, 'Au' : 196.967, 'Hg' : 200.592,\
             'Tl' : 204.383, 'Pb' : 207.2, 'Bi' : 208.980, 'Po' : 208.982,\
             'At' : 209.987, 'Rn' : 222.081, 'Fr' : 223.020, 'Ra' : 226.025,\
             'Ac' : 227.028, 'Th' : 232.038, 'Pa' : 231.036, 'U' : 238.029,\
             'Np' : 237, 'Pu' : 244, 'AM' : 243, 'Cm' : 247, 'Bk' : 247,\
             'Ct' : 251, 'Es' : 252, 'FM' : 257, 'Md' : 258, 'No' : 259,\
             'Lr' : 262, 'Rf' : 261, 'DB' : 262, 'Sg' : 266, 'Bh' : 264,\
             'Hs' : 269, 'Mt' : 268, 'DS' : 271, 'Rg' : 272, 'Cn' : 285,\
             'Nh' : 284, 'Fl' : 289, 'MC' : 288, 'Lv' : 292, 'Ts' : 294,\
             'Og' : 294}
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-be0db4c3921d3a6c, points: 0, grade: true, solution: false -->
```python
O2 = "O2"
assert abs(molekylmassa(O2) - 31.998) < 1.0e-4
NaCl = "NaCl"
assert abs(molekylmassa(NaCl) - 58.443) < 1.0e-4
glucose = "C6H12O6"
assert abs(molekylmassa(glucose) - 180.156) < 1.0e-4
hemoglobin = "C2952H4464N3248O812S8Fe4"
assert abs(molekylmassa(hemoglobin) - 98922.016) < 1.0e-4
```

# Stökiometribestämning: mass edition!

Kombinera föregående 2 övningar för att beräkna slutliga mängd i gram från reaktant mängder i gram.

<!-- grade_id: cell-158c189cb48293fe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ec692a1ef8244e6f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```