```python
import matplotlib.pyplot as plt
import numpy as np
```

# Titrerkurva

Titrerkurvan gör det möjligt att bestämma pKa för en syra-baspar. För en enkel syra-baspar bör titrerkurvan vara en sigmoid kurva och dess inflektionspunkter (förändring av andra derivator eller extrempunkter för första derivator) är relaterade till pKa för paret.

Vi läser in den experimentella titrerkurvans data från https://www.mathworks.com/matlabcentral/answers/456090-fitting-model-to-titration-data

```python
import csv

data=[]
with open('Titration.csv') as csvfile:
    spamreader = csv.reader(csvfile, delimiter=';')
    for row in spamreader:
        data.append(row)
volume = np.array(data[0][1:]).astype(float)
pH = np.array(data[1][1:]).astype(float)

plt.figure()
plt.plot(volume, pH)
plt.xlabel("Titration volume")
plt.ylabel("Measured pH")
plt.show()
```

Från detta kan du beräkna den första derivatan numeriskt genom att använda uttrycket:

$$ f'(x) \approx \frac{f(x+h) - f(x)}{h}$$

**tips:** använd slices!

<!-- grade_id: cell-5113a1ce6a9b470d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Från derivatan kan du nu hitta ekvivalenspunkten som det maximala värdet i derivatan. Sedan är pH vid halv ekvivalenspunkt pKa för ditt syra-bas-par:

<!-- grade_id: cell-81418c271af32de5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e9226def11562d84, points: 0, grade: true, solution: false -->
```python
assert abs(equivalence - 4.8)<=0.1
assert abs(pka - 4.44)<0.1
```

Observera att skillnaden mellan intilliggande punkter också kan göras med hjälp av den inbyggda funktionen numpy.diff för likvärdiga resultat.

```python
deriv = np.diff(pH)/np.diff(volume)

plt.figure()
plt.plot(volume[1:], deriv)
plt.xlabel("Titration volume (mL)")
plt.ylabel("pH derivative")
plt.show()
```

# Hastighetskonstanter

Enzymkinetik kan beskrivas av Michaelis-Menten kinetikmodellen.

$$ E + S  \underset{k_{-1}}{\overset{k_1}{\rightleftarrows}} ES \overset{k_2}{\rightarrow} E + P $$

Enzymet och substratet binder först i en reversibel reaktion och sedan sker reaktionen där produkten bildas och enzymet förblir intakt.

Detta ger upphov till en uppsättning differentialekvationer:

$$ \frac{d[S]}{dt} = -k_1 [E][S] + k_{-1} [ES]$$
$$ \frac{d[E]}{dt} = -k_1 [E][S] + (k_{-1} + k_2) [ES]$$
$$ \frac{d[ES]}{dt} = k_1 [E][S] - (k_{-1} + k_2) [ES]$$
$$ \frac{d[P]}{dt} = k_2 [ES] $$

Låt oss lösa dessa differentialekvationer med Euler-metoden, vilket innebär att vi tar ett litet tidssteg och löser $\frac{dA}{dt} = B$ som $A(t+dt) = A(t) + B * dt$.

Använd en numpy-array av storlek n_steps för varje kemiskt ämne, loopa över tidssteg och sprida ekvationerna. Plotta koncentrationen av varje ämne över tiden (du kan plotta E och ES separat eftersom deras koncentration vanligtvis är låg).

<!-- grade_id: cell-7d839a3d1b4904b1, points: 0, grade: false, solution: true -->
```python
dt = 0.01 # 0.1 s
k1 = 10**2
k_1 = 0.5
k2 = 40
initial_concentrations = [0.1, 10**(-4), 0, 0] # [S], [E], [ES], [P]
nsteps = 20000

# YOUR CODE HERE
raise NotImplementedError()
```

# Balansera kemiska ekvationer (svårt)

Vi har en utmanande kemisk ekvation att balansera.

K$_4$Fe(SCN)$_6$ + K$_2$Cr$_2$O$_7$ + H$_2$SO$_4$ →  Fe$_2$(SO$_4$)$_3$ + Cr$_2$(SO$_4$)$_3$ + CO$_2$ + H$_2$O + K$_2$SO$_4$ + KNO$_3$

Låt oss skapa ett generellt program för att göra detta med hjälp av numpy.

Först behöver vi skapa en lista av atomer och hur många det finns. Vi hade skapat en liknande funktion när vi beräknade molekylmassa:

```python
def molekylmassa(string):
    n_chars = len(string)
    total_mass = 0
    for i, atom in enumerate(string):
        if atom.isupper():
            j = i+1
            if j < n_chars and string[j].islower():
                atom += string[j]
                j += 1
            
            mass = molmassa[atom]
            # Get the number of atoms
            number = ""
            while j < n_chars and string[j].isnumeric():
                number += string[j]
                j += 1
            if len(number)>0:
                number = int(number)
            else:
                number = 1
            total_mass += number * mass
    return total_mass
```

**Uppgift 1:** Modifiera den här funktion så att det skapar en lexikon med atomer och deras antal.

<!-- grade_id: cell-198732c82f4edf2c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2c8b7ab00a872825, points: 0, grade: true, solution: false -->
```python
assert get_atoms("CO2") == {'C': 1, 'O': 2}
assert get_atoms("K2Cr2O7") == {'K': 2, 'Cr': 2, 'O': 7}
```

**Uppgift 2:** Modifiera den ytterligare så att det funkar med parentes (t.ex. Fe(CO)6).

<!-- grade_id: cell-1445ce2be4341340, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fafb5d3189db6a8c, points: 0, grade: true, solution: false -->
```python
assert get_atoms("Fe(CO)6") == {'Fe': 1, 'C': 6, 'O': 6}
assert get_atoms("Fe(CO)3(CN)3") == {'Fe': 1, 'C': 6, 'O': 3, 'N':3}
```

**Uppgift 3**: Använd denna funktionen för att kontrollera att båda sidor av ekvationen har samma grundämne.

**tips:** du kan använda `split`

<!-- grade_id: cell-d44258099e83191b, points: 0, grade: false, solution: true -->
```python
equation = "K4Fe(SCN)6 + K2Cr2O7 + H2SO4 =  Fe2(SO4)3 + Cr2(SO4)3 + CO2 + H2O + K2SO4 + KNO3"

# YOUR CODE HERE
raise NotImplementedError()

assert set(left_atoms) == set(right_atoms)
```

**Uppgift 3:**

Nu kommer vi att beskriva varje molekyl i ekvationen som en array med storlek N, där N är det totala antalet unika atomer som hittades. Varje tal i denna array kommer att motsvara antalet av den specifika atomen som förekommer i den givna molekylen.

Till exempel, om vi har atomerna C, O, N, H, Fe, så kommer:

if we have atoms C, O, N, H, Fe, then:
* NH3 att representeras som [0, 0, 1, 3, 0]
* H2O kommer att representeras som [0, 1, 0, 2, 0]
* Fe(CO)6 kommer att representeras som [6, 6, 0, 0, 1]

Det finns många sätt att göra detta, men här är ett förslag:
* Dela upp ekvationen i enskilda molekyler och loopa över dessa molekyler.
* Skapa för varje molekyl en noll-array med rätt storlek.
* Använd get_atoms för att hitta atomerna och koefficienterna.
* Lägg till antalet i arrayen på en position som beror på atomen (du kan använda `left_atoms.index(atom)`)
* Lagra arrayerna för varje molekyl i en lista.

<!-- grade_id: cell-98d5dae3cbfad6a1, points: 0, grade: false, solution: true -->
```python
N = len(left_atoms)
molecule_list = []
left, right = equation.split("=")
for molecule in left.split("+") + right.split("+"):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Efter att ha gjort detta har vi faktiskt bildat ett linjärt ekvationssystem:

$$ x_0 A_0 + x_1 A_1 + ... = 0 $$

Detta är verkligen det vi vill ha, under antagandet att vi har skrivit `vänster = höger`-kemisk ekvation som `vänster - höger = 0`. Detta kan också ses som en matrisekvation AX = 0.

Det finns givetvis ett oändligt antal lösningar till detta problem som är multiplar av varandra. Så istället väljer vi att $x_0 = 1$. För detta flyttar vi helt enkelt det till höger och löser.

$$ x_1 A_1 + ... = -A_0  $$

Det ger oss exakt 8 ekvationer och 8 okända. Vi kan nu använda numpy.linalg.solve som är avsett att lösa $AX = B$ matrisekvationer.

Obs: Om vi inte hade haft samma antal ekvationer och okända hade vi kunnat använda numpy.linalg.lstsq som försöker hitta den bästa lösningen på systemet.

```python
A = np.array(molecule_list).T
solution = np.linalg.solve(A[:,1:],-A[:,0])
print(solution)
```

Vi ser att de sista siffrorna är negativa, men det är normalt eftersom vi löser vänster - höger = 0. När vi återgår till höger kommer de att bli positiva igen.

Men ett annat olyckligt problem är att vår lösning inte är ett heltal. Det är fortfarande en giltig lösning, men vi föredrar en heltalslösning. Vi märker att det finns många "0.166666667" och därför har vi goda förhoppningar om att multiplicera med 6 kommer att ge heltalsresultat (detta motsvarar att sätta $c0 = 6$).

```python
print(solution*6)
```

Detta ger oss vår slutlig ekvation

6 K$_4$Fe(SCN)$_6$ + 97 K$_2$Cr$_2$O$_7$ + 355 H$_2$SO$_4$ →  3 Fe$_2$(SO$_4$)$_3$ + 97 Cr$_2$(SO$_4$)$_3$ + 36 CO$_2$ + 355 H$_2$O + 91 K$_2$SO$_4$ + 36 KNO$_3$

som du kan kontrollera är balanserad!

```python

```